# Homework 6: NumPy(29 pts)

name: Kennon Stewart

email: ksstewar@umich.edu

This homework assignment took me XX hours in total to complete. (Please help us to gauge the difficulty of the assignment.)

## Collaboration Disclosure

In the cell below, please list *everyone* with whom you discussed any of the homework problems, excluding only the GSIs and the course instructor. 

If you did not discuss the homework with anyone else, write __"I did not discuss this homework with anyone."__

Even if you discuss questions with other, the code you submit must be only yours. All work is checked with the [MOSS plagiarism detector](https://theory.stanford.edu/~aiken/moss/).

---

## Submission Instructions
Your homework solutions should be written entirely in this Jupyter notebook file. Once it contains your solutions, you should submit this notebook through Canvas. 


Before submitting, please make sure to __Cells->Run All__ executes without errors; errors in your code translate directly to point deductions. 
In general, you don't need to do explicitly raise errors (e.g. with the ```raise``` function) if we don't ask you to in the problem statement.
However, even in cases where we ask you to check for errors, your submission should not contain any examples of your functions actually raising those errors.

Note that many parts of this homework where you are expected to type in code will have ```NotImplementedError()``` as a placeholder. You need to delete this function and replace it with your own code.

## Homework tips 

1. **Start early!** If you run into trouble installing things or importing packages, it’s
best to find those problems well in advance, not the night before your assignment is
due when we cannot help you!

2. **Make sure you back up your work!** At a minimum, do your work in a Dropbox
folder. Better yet, use git, which is well worth your time and effort to learn.

3. **Be careful to follow directions!** Remember that Python is case sensitive. If
you are ask you to define a function called my_function and you define a function
called My_Function, you will not receive full credit. You may want to copy-paste
the function names below to make sure that the functions in your notebook match.

## Error checking

You do not need to do error checking (raising errors, etc.) in your code unless we explicitly ask you to so in a problem.


## Nbgrader

We will be using `nbgrader` to grade your jupyter notebook. You will notice some `read-only` cells in the assignment that contain `assert` statements. These are tests that your code must pass for your solution to be correct. If any of the tests fail, you will get an python error and not get points for that question. 

**Note:** The tests shown not are not comprehensive; additional tests will be used at grading time. You are encouraged to read the problem carefully and verify your code covers all possible cases.

**Be careful:** If a jupyter notebook cell takes longer than `60s` to run, the autograder will not grade it and you will receive zero credit for that question.

## Question 1:. The Bakersville Bake-off! (3 pts)

The houses in the town of Bakersville are arranged on a grid with R rows (numbered 1 through R) and C columns (numbered 1 through C). The house located on row r and on column c is denoted by (r, c). We say two houses are **neighbors** if they have a common side. 

During the annual Bakersville Bake-off, each household bakes a certain number of cakes to share with their neighbors.
The town comptroller keeps track in an $R \times C$ matrix $A$, where $A_{r, c}$ is the number of cakes which household $(r, c)$ has baked.

The stability of the social order in Bakersville depends on each house baking enough cakes so they can give at least one cake to each neighbor and submit one cake to be judged. Thus, we say a house is **stable** if the number of cakes $A_{r, c}$ is strictly greater than its number of neigbors. We say Bakersville is **stable** if every house is stable. 

The comptroller has asked you to define a function `is_stable` that takes an $R \times C$ numpy array (i.e., an ndarray) as its only argument and returns a Python Boolean that is True if and only if the argument is a stable array and False otherwise. 

<!-- grade_id: is_stable, points: 0, grade: false, solution: true -->
```python
import numpy as np
def is_stable(input_array):
    
    # initialize variables
    input_array = np.array(input_array)
    print(f"shape of the input matrix: {input_array.shape}")
    try:
        r, c = input_array.shape
    except ValueError:
        dimension = len(input_array)
        print(f"single-row matrix shape: {dimension}")
        neighbors = 0
        
        for house_index in range(dimension):
            # check for right neighbor
            if house_index > 0:
                neighbors += 1

            # check for left neighbor
            if house_index < dimension - 1:
                neighbors += 1

            # compare against neighbors
            if input_array[house_index] < neighbors + 1:
                return False
    
    # parse matrix with false condition ready
    for street_index in range(r):
        for house_index in range(c):
            neighbors = 0
            # check for a neighbor to the left
            if house_index > 0:
                neighbors += 1
            
            # check for a neighbor to the right
            if house_index < c - 1:
                neighbors += 1
            
            # check for a neighbor on top
            if street_index > 0:
                neighbors += 1
            
            # check for a neighbor below
            if street_index < r - 1:
                neighbors +=1 
            print(f"number of neighbors for house ({street_index}, {house_index}): {neighbors}")
            
            # if value is greater than number of neighbors + 1
            cookies = input_array[street_index][house_index]
            if cookies < neighbors + 1:
                return False
    
    # return true when parsing complete
    return True
```

<!-- grade_id: correct_stable, points: 3, grade: true, solution: false -->
```python
array1 = np.array([4, 7, 8,
                   5, 10, 9,
                   3, 5, 7]).reshape(3, 3)
array2 = np.array([3, 2, 4,
                   1, 1, 2,
                   3, 3, 3]).reshape(3, 3)


assert is_stable(array1)
assert not is_stable(array2)
assert not is_stable(np.array([1., 0., 1.]))
```

Output:
```
shape of the input matrix: (3, 3)
number of neighbors for house (0, 0): 2
number of neighbors for house (0, 1): 3
number of neighbors for house (0, 2): 2
number of neighbors for house (1, 0): 3
number of neighbors for house (1, 1): 4
number of neighbors for house (1, 2): 3
number of neighbors for house (2, 0): 2
number of neighbors for house (2, 1): 3
number of neighbors for house (2, 2): 2
shape of the input matrix: (3, 3)
number of neighbors for house (0, 0): 2
number of neighbors for house (0, 1): 3
shape of the input matrix: (3,)
single-row matrix shape: 3
```

## Question 2: Down with `for` loops (12 pts)
Each of the problems below contains a function that uses `for` loops to perform a certain operation on arrays. Your job is to re-write these function using only Numpy array manipulations and library functions (e.g. `np.xxx()`). Do not use any `for` or `while` loops, iterators, generators, or list comprehensions in your solutions.

**2(a)**  Return all of the rows of the 2-D integer matrix `A` where where each entry of the row is distinct (3 pts)

```python
def distinct_rows_py(A):
    'Return all rows of A that have completely distinct entries.'
    return np.array([a for a in A if len(set(a)) == len(a)])
A = np.eye(5)
print(distinct_rows_py(A))
B = np.arange(9).reshape(3, 3)
print(distinct_rows_py(B))
```

Output:
```
[]
[[0 1 2]
 [3 4 5]
 [6 7 8]]
```

<!-- grade_id: distinct_rows_np, points: 0, grade: false, solution: true -->
```python
def distinct_rows_np(A):
    # rule: the returned row should be equal in length
    # to the set of elements
    filtered_matrix = A[np.apply_along_axis(lambda x: True if len(x) == len(set(x)) else False, 1, A)]
    return filtered_matrix
    
# run the test cases
geek_array = np.array([[8,1,8],
                         [4,3,9],
                         [5,2,6],
                         [8,1,8]])
distinct_rows_np(geek_array)
```

Output:
```
array([[4, 3, 9],
       [5, 2, 6]])
```

<!-- grade_id: correct_distinct, points: 3, grade: true, solution: false -->
```python
A = np.eye(5)
np.testing.assert_allclose(distinct_rows_np(A), np.empty([0, 5]))
```

**2(b)** Given a vector $v$ of length $n$, and an integer $0<k<=n$, return a 2D array as: (3 pts)
```
[[v[0], ..., v[k-1]],
 [v[1], ..., v[k]  ],
 [v[2], ..., v[k+1]],
 [ .     .       . ],
 [ .     .       . ],
 [ .     .       . ],
 [v[n-k+1, ..., v[n]]
 ```

```python
def return_window_indices(input_array, index, k):
    return input(array_index[index:index+k])
```

```python
def sliding_stack_py(v, k):
    # found some random function on NumPy's website
    # not sure if this counts as an iterable loop
    output = np.apply_along_axis(lambda x: return_window_indices(x), 1, A)
```

<!-- grade_id: sliding, points: 0, grade: false, solution: true -->
```python
import numpy as np
def sliding_stack_np(input_array, window_size):
    window_shape = len(input_array) - window_size + 1
    # breaking down the method to keep better track of the keyword arguments
    return np.lib.stride_tricks.as_strided(input_array, 
                                           shape=(window_shape, window_size), 
                                           strides=(input_array.strides[0], 
                                           input_array.strides[0]))
```

<!-- grade_id: correct_sliding, points: 3, grade: true, solution: false -->
```python
np.testing.assert_allclose(sliding_stack_np(np.array([1, 2, 3, 4, 5]), 3), np.array([[1, 2, 3],[2, 3, 4],[3, 4, 5]]))
np.testing.assert_allclose(sliding_stack_np(np.array([1, 2, 3, 4, 5]), 5), np.array([[1, 2, 3, 4, 5]]))
```

**2(c)** Given a vector of non-negative integers `v`, with `max(v) = m`, return a vector `c` of length `m + 1` such that `c[i]` is the number of times that the integer `i` appears in `v`.(3 pts)

```python
def digit_count_py(v):
    m = max(v)
    ret = np.zeros(m + 1, int)
    for vv in v:
        ret[vv] += 1
    return ret
v = np.array([0, 0, 1, 1, 2, 2, 2, 3, 5])
digit_count_py(v)
```

Output:
```
array([2, 2, 3, 1, 0, 1])
```

<!-- grade_id: digit_count_np, points: 0, grade: false, solution: true -->
```python
import numpy as np
def digit_count_np(v):
    return np.bincount(v)
```

<!-- grade_id: correct_digit_count_np, points: 3, grade: true, solution: false -->
```python
v = np.array([0, 0, 1, 1, 2, 2, 2, 3, 4])
np.testing.assert_allclose(digit_count_np(v), np.array([2, 2, 3, 1, 1]))
v = np.array([0, 0, 0, 0, 2, 2, 2, 3, 4])
np.testing.assert_allclose(digit_count_np(v), np.array([4, 0, 3, 1, 1]))
```

**2(d)** Call a square $n\times n$ matrix $A$ *countersymmetric* if $A_{ij} = A_{n-j,n-i}$ for all $i$ and $j$. An example of such a matrix is:

$$
\begin{pmatrix}
4 & 3 & 2 & 1 & 0\\
8 & 7 & 6 & 5 & 1\\
11 & 10 & 9 & 6 & 2\\
13 & 12 & 10 & 7 & 3\\
14 & 13 & 11 & 8 & 4
\end{pmatrix}
$$

Write a function `is_countersym` that checks this property: (3 pts)

```python
def is_countersym_py(A):
    "Returns True if A is countersymmetric"
    n = A.shape[0]
    for i in range(n):
        for j in range(n):
            if A[i, j] != A[n - j - 1, n - i - 1]:
                return False
    return True
cs_matrix = np.array([[ 4,  3,  2,  1,  0], [ 8,  7,  6,  5,  1], [11, 10,  9,  6,  2], [13, 12, 10,  7,  3], [14, 13, 11,  8,  4]])
is_countersym_py(cs_matrix)
```

Output:
```
True
```

<!-- grade_id: is_countersym_np, points: 0, grade: false, solution: true -->
```python
import numpy as np
def is_countersym_np(A):
    # I couldn't find the exact method that enables countersymmetry
    # but flipping the matrix twice seems to do the trick?
    A_transpose = np.matrix(A).transpose().transpose()
    print(f"original matrix:\n{A}\ntranspose matrix:\n{A_transpose}")
    
    if np.array_equal(A, A_transpose):
        return True
    return False
```

<!-- grade_id: correct_is_countersym_np, points: 3, grade: true, solution: false -->
```python
cs_matrix = np.array([[ 4,  3,  2,  1,  0], [ 8,  7,  6,  5,  1], [11, 10,  9,  6,  2], [13, 12, 10,  7,  3], [14, 13, 11,  8,  4]])
assert is_countersym_np(cs_matrix), "The matrix is not counter-symmetric"

```

Output:
```
original matrix:
[[ 4  3  2  1  0]
 [ 8  7  6  5  1]
 [11 10  9  6  2]
 [13 12 10  7  3]
 [14 13 11  8  4]]
transpose matrix:
[[ 4  3  2  1  0]
 [ 8  7  6  5  1]
 [11 10  9  6  2]
 [13 12 10  7  3]
 [14 13 11  8  4]]
```

## Question 3: $k$-means clustering (14 pts)

$k$-means is a fundamental algorithm for clustering multivariate data. The inputs to the algorithm are:
- An $n\times p$ data matrix $X$ consisting of $n$ observations of a $p$-dimensional feature vector, and
- A $k\times p$ matrix $C$ containing initial guesses for each $k$ cluster centers.

The algorithm proceeds by iteratively a) assigning each point to the nearest cluster center, and b) recomputing the cluster centers as the mean of all of the currently assigned points. Here is a partial implementation:

```python
def kmeans(X, C):
    """
    K-means algorithm.

    Args:
        - X: ndarray, shape (n, p), n observations of p-dimensional feature vector
        - C: ndarray, shape (k, p), k initial cluster centers

    Returns:
        Tuple of length two:
        The first entry is integer ndarray, shape (n), cluster assignments for each data point
        The second entry is ndarray, shape (k, p), centers of each cluster
    """
    assert X.shape[1] == C.shape[1]  # p should match
    old_assignments = None
    while True:
        assignments = nearest_cluster(X, C)
        if np.all(assignments == old_assignments):
            # converged
            return assignments, C
        old_assignments = assignments
        C = compute_centroids(X, assignments)
```

You will finish implementing this algorithm by completing the missing functions `nearest_cluster()` and `compute_centroids()` below. Note: as in the Question 1, avoid using loops, iterators, or comprehensions. You can find solutions with only broadcasting

**3(a)** Implement the function `nearest_cluster`. It should take two array arguments, the data points `X` and the cluster centers `C`, and return an integer array giving the index in `C` which is nearest to each point in `X`. (3 pts)

<!-- grade_id: 3a, points: 0, grade: false, solution: true -->
```python
def nearest_cluster(X, C):
    """
    For each point in X, find the nearest point in C.

    Args:
        X: ndarray, shape (n, p), n points of dimension p.
        C: ndarray, shape (k, p), k points of dimension p.

    Returns:
        Integer array of length n, [j[1], j[2], ..., j[n]], such that |X[i] - C[j[i]]| <= |X[i] - C[ell]| for 1 <= ell <= k.
    """
    # calculate the positive distance between each input point X
    # and the existing mean clusters in C
    print(f"input X:\n{X}\ninput C:\n{C}")
    adjusted_input_matrix = abs(X[:, np.newaxis] - C)
    print(f"distance between each of the points:\n{adjusted_input_matrix}")
    sum_of_distances = np.sum(adjusted_input_matrix, axis = 2)
    print(f"reshaped matrix:\n{sum_of_distances}")
    
    # use the distances to determine the closest point in C
    nearest_indices = np.argmin(sum_of_distances, axis=1)
    print(f"find the index of the smallest distance:\n{nearest_indices}")
    return nearest_indices

ans = np.array([0, 0, 0, 0, 1, 1, 1, 2, 2, 2])
X = np.array([[-1], [0.5], [1], [1.3], [1.6], [2], [2.3], [2.7], [3], [100]])
C = np.array([[1], [2], [3]])
nearest_cluster(X, C)
```

Output:
```
input X:
[[ -1. ]
 [  0.5]
 [  1. ]
 [  1.3]
 [  1.6]
 [  2. ]
 [  2.3]
 [  2.7]
 [  3. ]
 [100. ]]
input C:
[[1]
 [2]
 [3]]
distance between each of the points:
[[[ 2. ]
  [ 3. ]
  [ 4. ]]

 [[ 0.5]
  [ 1.5]
  [ 2.5]]

 [[ 0. ]
  [ 1. ]
  [ 2. ]]

 [[ 0.3]
  [ 0.7]
  [ 1.7]]

 [[ 0.6]
  [ 0.4]
  [ 1.4]]

 [[ 1. ]
  [ 0. ]
  [ 1. ]]

 [[ 1.3]
  [ 0.3]
  [ 0.7]]

 [[ 1.7]
  [ 0.7]
  [ 0.3]]

 [[ 2. ]
  [ 1. ]
  [ 0. ]]

 [[99. ]
  [98. ]
  [97. ]]]
reshaped matrix:
[[ 2.   3.   4. ]
 [ 0.5  1.5  2.5]
 [ 0.   1.   2. ]
 [ 0.3  0.7  1.7]
 [ 0.6  0.4  1.4]
 [ 1.   0.   1. ]
 [ 1.3  0.3  0.7]
 [ 1.7  0.7  0.3]
 [ 2.   1.   0. ]
 [99.  98.  97. ]]
find the index of the smallest distance:
[0 0 0 0 1 1 1 2 2 2]
array([0, 0, 0, 0, 1, 1, 1, 2, 2, 2])
```

<!-- grade_id: correct_3a, points: 3, grade: true, solution: false -->
```python
ans = np.array([0, 0, 0, 0, 1, 1, 1, 2, 2, 2])
X = np.array([[-1], [0.5], [1], [1.3], [1.6], [2], [2.3], [2.7], [3], [100]])
C = np.array([[1], [2], [3]])
np.testing.assert_allclose(nearest_cluster(X, C), ans)

X = np.array([[-101, 1], [1,1], [1, 1.01], [1.5, 1.6], [2.5, -3], [2.5, 2.7], [3, 33]])
C = np.array([[1, 1], [2,2], [3,3]])
ans = np.array([0, 0, 0, 1, 0, 2, 2])
np.testing.assert_allclose(nearest_cluster(X, C), ans)
```

Output:
```
input X:
[[ -1. ]
 [  0.5]
 [  1. ]
 [  1.3]
 [  1.6]
 [  2. ]
 [  2.3]
 [  2.7]
 [  3. ]
 [100. ]]
input C:
[[1]
 [2]
 [3]]
distance between each of the points:
[[[ 2. ]
  [ 3. ]
  [ 4. ]]

 [[ 0.5]
  [ 1.5]
  [ 2.5]]

 [[ 0. ]
  [ 1. ]
  [ 2. ]]

 [[ 0.3]
  [ 0.7]
  [ 1.7]]

 [[ 0.6]
  [ 0.4]
  [ 1.4]]

 [[ 1. ]
  [ 0. ]
  [ 1. ]]

 [[ 1.3]
  [ 0.3]
  [ 0.7]]

 [[ 1.7]
  [ 0.7]
  [ 0.3]]

 [[ 2. ]
  [ 1. ]
  [ 0. ]]

 [[99. ]
  [98. ]
  [97. ]]]
reshaped matrix:
[[ 2.   3.   4. ]
 [ 0.5  1.5  2.5]
 [ 0.   1.   2. ]
 [ 0.3  0.7  1.7]
 [ 0.6  0.4  1.4]
 [ 1.   0.   1. ]
 [ 1.3  0.3  0.7]
 [ 1.7  0.7  0.3]
 [ 2.   1.   0. ]
 [99.  98.  97. ]]
find the index of the smallest distance:
[0 0 0 0 1 1 1 2 2 2]
input X:
[[-101.      1.  ]
 [   1.      1.  ]
 [   1.      1.01]
 [   1.5     1.6 ]
 [   2.5    -3.  ]
 [   2.5     2.7 ]
 [   3.     33.  ]]
input C:
[[1 1]
 [2 2]
 [3 3]]
distance between each of the points:
[[[1.02e+02 0.00e+00]
  [1.03e+02 1.00e+00]
  [1.04e+02 2.00e+00]]

 [[0.00e
```

**3(b)** Implement the function `compute_centroids`. It should take two array arguments, the data points `X` and the assignment array `a`, and return an $k \times p$ array containing the cluster centroids (averages) for each point assigned to cluster $0, \dots, k-1$. (You may assume that every entry of $a$ is between $0$ and $k-1$, inclusive.) (3 pts) 

<!-- grade_id: 3b, points: 0, grade: false, solution: true -->
```python
def compute_centroids(X, a):
    # what is the number of clusters?
    print(X)
    cluster_count = max(a)+1  # +1 because cluster labels start from 0

    # create a list where every element is a cluster index
    cluster_indices = np.arange(len(X))
    print(f"indices of the cluster:\n{cluster_indices}")
    
    # Use np.add.at to accumulate sums for each cluster
    sums = np.zeros((cluster_count, X.shape[1]))
    counts = np.zeros(cluster_count)

    # Accumulate sums and counts using advanced indexing
    np.add.at(sums, a, X)  # Sum points in each cluster
    np.add.at(counts, a, 1)  # Count points in each cluster
    print(f"sums of the clusters:\n{sums}\ncounts of the clusters:\n{counts})")

    # Compute centroids, avoiding division by zero
    centroids = np.divide(sums, counts[:, np.newaxis], out=np.zeros_like(sums), where=counts[:, np.newaxis] != 0)

    return centroids
```

<!-- grade_id: correct_3b, points: 3, grade: true, solution: false -->
```python
import numpy as np
np.random.seed(1)
X = np.random.rand(30, 2)
a = np.random.randint(0, 5, size=30)
assert compute_centroids(X, a).shape == (5, 2)
assert np.allclose(X[a == 0].mean(axis=0), compute_centroids(X, a)[0])
```

Output:
```
[[4.17022005e-01 7.20324493e-01]
 [1.14374817e-04 3.02332573e-01]
 [1.46755891e-01 9.23385948e-02]
 [1.86260211e-01 3.45560727e-01]
 [3.96767474e-01 5.38816734e-01]
 [4.19194514e-01 6.85219500e-01]
 [2.04452250e-01 8.78117436e-01]
 [2.73875932e-02 6.70467510e-01]
 [4.17304802e-01 5.58689828e-01]
 [1.40386939e-01 1.98101489e-01]
 [8.00744569e-01 9.68261576e-01]
 [3.13424178e-01 6.92322616e-01]
 [8.76389152e-01 8.94606664e-01]
 [8.50442114e-02 3.90547832e-02]
 [1.69830420e-01 8.78142503e-01]
 [9.83468338e-02 4.21107625e-01]
 [9.57889530e-01 5.33165285e-01]
 [6.91877114e-01 3.15515631e-01]
 [6.86500928e-01 8.34625672e-01]
 [1.82882773e-02 7.50144315e-01]
 [9.88861089e-01 7.48165654e-01]
 [2.80443992e-01 7.89279328e-01]
 [1.03226007e-01 4.47893526e-01]
 [9.08595503e-01 2.93614148e-01]
 [2.87775339e-01 1.30028572e-01]
 [1.93669579e-02 6.78835533e-01]
 [2.11628116e-01 2.65546659e-01]
 [4.91573159e-01 5.33625451e-02]
 [5.74117605e-01 1.46728575e-01]
 [5.89305537e-01 6.99758360e-01]]
indices o
```

**3(c)** The performance of the $k$-means algorithm is known to depend heavily on the starting point (the initial clusters `C` passed in as the second argument.) In some cases, using a "good" starting point can dramatically improve the performance of the algorithm.

The $k$-means++ algorithm is designed to find such a good starting point. [According to Wikipedia](https://en.wikipedia.org/wiki/K-means%2B%2B), the steps of $k$-means++ are:

1. Choose one center uniformly at random among the data points.
2. For each data point $x$ not chosen yet, compute $D(x)$, the distance between $x$ and the nearest center that has already been chosen.
3. Choose one new data point at random as a new center, using a weighted probability distribution where a point $x$ is chosen with probability proportional to $D(x)^2$.
4. Repeat Steps 2 and 3 until $k$ centers have been chosen.

Implement this algorithm using the skeleton provided below. As before, your implementation should only use Numpy functions--no additional loops or comprehensions. 

**Note**: To ensure reproducibility, the parts of the algorithm that rely on ranndomness are provided for you. Your job is to fill in the missing lines necessary to complete the algoritm. (5 pts.)

<!-- grade_id: 3c, points: 0, grade: false, solution: true -->
```python
def kmeanspp(X, k, rng):
    """
    k-means++ algorithm.

    Args:
        - X: ndarray, shape (n, p), as above.
        - k, the number of clusters.
        - rng: instance of np.random.Generator().

    Returns:
        ndarray, shape (k, p), cluster centers.
    """
    n, p = X.shape
    C = np.zeros((k, p))
    
    # Step 1: Choose the first center randomly
    j = rng.choice(n)
    C[0] = X[j]
    
    # Step 2 and 3: Choose remaining centers
    for i in range(1, k):
        # Compute distances from each point to the nearest center
        D = np.linalg.norm(X[:, np.newaxis] - C[:i], axis=2)  # shape (n, i)
        min_distances = np.min(D, axis=1)  # shape (n,)
        
        # Calculate weights based on squared distances
        w = min_distances ** 2
        
        # Normalize the weights to sum to 1
        w /= np.sum(w)
        
        # Choose the next center based on the weighted probabilities
        j = rng.choice(n, p=w)
        C[i] = X[j]
        
    return C
```

<!-- grade_id: correct_3c, points: 5, grade: true, solution: false -->
```python
import numpy as np

np.random.seed(1)
X = np.random.rand(30, 4)
k = 2
rng = np.random.default_rng(12345)
assert kmeanspp(X, k, rng).shape == (2, 4)
ans = np.array([[0.87638915, 0.89460666, 0.08504421, 0.03905478], 
                [0.6634415 , 0.62169572, 0.11474597, 0.94948926]])
np.testing.assert_allclose(kmeanspp(X, k, rng), ans)
```

**3(d)** In order to measure how good a clustering is, we can define the *within-class variance* 

$$ V(\mathbf{X}, \mathbf{a}, \mathbf{C}) = \sum_{i=1}^n \| \mathbf{x}_i - \mathbf{c}_{a_i} \|^2,$$

where the $i$-th element of $\mathbf{a}=\{a_1,\dots,a_n\}$ is the cluster assignment of observation $i$, and $\mathbf{C}=(\mathbf{c}_1,\dots,\mathbf{c}_k)$ are the centers of each cluster. Thus, $V(\mathbf{X}, \mathbf{a}, \mathbf{C})$ is the sum of the squared distance from each data point to the center of its assigned cluster.

Implement this function. (Again, no loops, just use Numpy functions.) (3 pts) 

<!-- grade_id: 3d, points: 0, grade: false, solution: true -->
```python
def V(X, a, C):
    # Get the number of data points
    n = len(X)
    
    # Select the centers based on the assignments
    centers = C[a]  # shape (n, p), each point is assigned to its respective center
    
    # Compute the squared distances
    squared_distances = np.sum((X - centers) ** 2, axis=1)  # shape (n,)
    
    # Return the sum of squared distances
    return sum(squared_distances)
```

<!-- grade_id: correct_3d, points: 3, grade: true, solution: false -->
```python
import numpy as np

X = np.arange(9).reshape(3, 3)
a = np.arange(3)
C = 2 + X
np.testing.assert_allclose(V(X, a, C), 36)

```